---
title: 'Starting Markdown'
description: 'Introduction to markdown.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Markdown is a system for writing text documents that is designed to be easier than html.  It is designed based on a more generic system called Markdown with some special modifications to make it work with R.

`@instructions`
Let's start by creating a totally blank markdown file called myfirstfile.rmd.    Do this by entering file.edit("~/myfirstfile.rmd")  in the console.

`@hint`
Type the code exactly as given

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
file.edit("~/myfirstfile.rmd")
```

`@solution`
```{r}
file.edit("~/myfirstfile.rmd")
```

`@sct`
```{r}

```
